```python
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np
```

# Let's start by loading a Miriam data file

```python
df = pd.read_csv('miriam_20202011_151410_a.csv')
```

```python
df
```

Output:
```
Time      U      M  extra    box  A1_neg  A2_neg  A3_neg  \
0   2020-11-20T14:12:01  62.10  66.00  66.00  49.81     695     763     701   
1   2020-11-20T14:13:02  62.28  64.98  58.42  49.55    1019     742     706   
2   2020-11-20T14:14:03  63.80  64.39  60.30  49.55    1022     718     709   
3   2020-11-20T14:15:04  65.39  64.58  60.47  49.68    1022     690     712   
4   2020-11-20T14:16:05  65.39  64.98  60.65  49.68    1023     671     715   
..                  ...    ...    ...    ...    ...     ...     ...     ...   
56  2020-11-20T15:08:57  65.39  65.59  61.01  50.48     966     972     872   
57  2020-11-20T15:09:58  65.39  65.39  61.01  50.48     967     971     874   
58  2020-11-20T15:10:59  65.39  65.39  61.01  50.48     968     972     875   
59  2020-11-20T15:12:00  65.39  65.59  61.01  50.48     979     973     877   
60  2020-11-20T15:13:01  65.39  65.59  61.01  50.48     991     971     879   

    A4_neg  A5_neg  ...  H4_pos23  H5_pos23  H6_pos23  H7_posctrl  \
0
```

If we inspect the column names, we can see that they have been modified to indicate what was loaded in each well. We note that only the column names with a `'_'` are the names of wells. The others describe time or temperature parameters. (These data are from an old version of Miriam with less understandable names for these parameters.) Therefore, we anticipate a little function like this will be useful:

```python
def has_underscore(colname):
    return '_' in colname
```

# Q1 Make a list of all column names called `colnames`

<!-- grade_id: cell-70920cb39b09b8d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Input In [4], in <cell line: 2>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-9ee6881020b3bdd3, points: 1, grade: true, solution: false -->
```python
assert len(colnames)==102
assert 'A1_neg' in colnames
assert 'Time' in colnames
```

# Q2 Let's plot all of our data to get an initial impression

Make a plot of the `Time` column against each of the other columns with fluorescence data. It should look like this.

![all_timeseries_chaos.png](all_timeseries_chaos.png)

<!-- grade_id: cell-2cf288da87927b03, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Note that each trace seems to start from a different level. Let's get rid of this.

# Q3 Now, let's subtract a "baseline" fluorescence from each well's data and replot. We do this by averaging the first 20 values from each well and then plotting as above.

The results should look like this

![baseline_subtracted.png](baseline_subtracted.png)

<!-- grade_id: cell-7191c919ec4ffeb0, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Q4 Finally, let's use our pre-existing knowledge of the contents of each well to color the trace.

We can set a variable `color` based on the column name like so:

```python
if 'neg' in colname:
    color = 'b'
elif 'pos' in colname:
    color = 'r'
else:
    color = 'k'
```

and then pass the `color` keyword argument to the `plot()` function. The results should look like this:

![colored-miriam-results.png](colored-miriam-results.png)

<!-- grade_id: cell-92d4461ce824bf15, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```