```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

SSH to your account at `wm01-tiger.strawlab.privat` and follow the puzzle starting in the directory `/schnitzeljagd`. Place the final answer in the variable `final_answer`.

<!-- grade_id: cell-d6b141313f4f3a09, points: 0, grade: false, solution: true -->
```python
final_answer = "laksdjf8alkds9339akklafk93"
```

<!-- grade_id: cell-3d5b7fe600c4c6ba, points: 1, grade: true, solution: false -->
```python
assert ads_hash(final_answer) == 'b03c9d8d1c'
```