```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

The questions here are in reference to the notebook at https://github.com/strawlab/pm21-ox/blob/main/lectures/lecture-13/02%20BEST.ipynb You will need to download that notebook and run it to answer the following questions.

```python
import numpy as np
```

## Problems

### Q1 Run the notebook above with the following samples. What is the mean value of the posterior estimate of the difference of means? Put this in the variable `posterior_diff_means`.

```
drug = (101,100,102,104,102,97,105,105,98,101,100,123,105,103,100,95,102,106,
        109,102,82,102,100,102,102,101,102,102,103,103,97,97,103,101,97,104,
        96,103,124,101,101,100,101,101,104,100,101)
placebo = (99,101,100,101,102,100,97,101,104,101,102,102,100,105,88,101,100,
           104,100,100,100,101,102,103,97,101,101,100,101,99,101,100,100,
           101,100,99,101,100,102,99,100,99)
```

Note: you may also try running this at [the online version of the BEST test](http://www.sumsar.net/best_online/).

<!-- grade_id: cell-4583b093ba19414e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-16393e9287358f37, points: 1, grade: true, solution: false -->
```python
# This is a test of the above, do not change this code.
assert(ads_hash(int(np.round(posterior_diff_means*10)))=='4a44dc1536')
```

### Q2 Note the bounds of the 94% highest posterior density interval (HDI) for the difference of means between group 1 and group 2. Read these values from the summary table of the trace looking at the "difference of means" variable.

Put these into the variables `diff_means_hdi_lower` and `diff_means_hdi_upper`

(94% was chosen by the authors of the software as a friendly reminder that this value is an arbitrary choice - rather than perhaps 95% which my not remind you of this choice.)

<!-- grade_id: cell-f10804d9e06925c6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c13982956b19759a, points: 1, grade: true, solution: false -->
```python
assert ads_hash(int(np.round(diff_means_hdi_lower*10)))=='d4735e3a26'
assert ads_hash(int(np.round(diff_means_hdi_upper+1.0))) == '4e07408562'
```

### Q3 Re-run the notebook with the following samples. What is the mean value for the posterior estimate of the difference of means for this smaller sample size? Put this in the variable `posterior_diff_means_small_samples`.¶

```
drug = (103,99,103,105,102,101)
placebo = (101,100,102,105,102,96)
```

<!-- grade_id: cell-e349ae3da9999c30, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-25d34717f657c572, points: 1, grade: true, solution: false -->
```python
# This is a test of the above, do not change this code.
assert(ads_hash(int(np.round(posterior_diff_means_small_samples*10)))=='4fc82b26ae')
```

### Q4 Note the bounds of the 94% highest posterior density interval (HDI) for the difference of means between group 1 and group 2 for this smaller data set. Read these values from the summary table of the trace looking at the "difference of means" variable.

Put these into the variables `diff_means_hdi_lower_small_samples` and `diff_means_hdi_upper_small_samples`


<!-- grade_id: cell-6cb965cd2f3f109a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b0b09eb0a16dfb2b, points: 1, grade: true, solution: false -->
```python
# This is a test of the above, do not change this code.
assert ads_hash(int(np.round(diff_means_hdi_lower_small_samples*2))) == '37aa1ccf80'
assert ads_hash(int(np.round(diff_means_hdi_upper_small_samples+1))) == 'e7f6c01177'
```

### Q5 The posterior estimate of the difference in means between the two datasets is similar. But one of the datasets has much more data than the other. Explain why the 94% HDI exclude a zero difference result in one dataset but not the other.

<!-- grade_id: cell-cde0c771165a691c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE